## Exercise 09.1 (checking data validity)

The Fibonacci series is valid only for $n \ge 0$. Add to the Fibonacci function in this notebook a check that raises an exception if $n < 0$. Try some invalid data cases to check that an exception is raised.

*Optional:* Use `pytest` to test that an exception *is* raised for some $n < 0$ cases.

<!-- grade_id: cell-982c3611466532a2, points: 0, grade: false, solution: true -->
```python
def f(n): 
    "Compute the nth Fibonacci number using recursion"
    if n == 0:
        return 0  # This doesn't call f, so it breaks out of the recursion loop
    elif n == 1:
        return 1  # This doesn't call f, so it breaks out of the recursion loop
    elif n < 0:
        raise ValueError("Cannot generate series.")
    else:
        return f(n - 1) + f(n - 2)  # This calls f for n-1 and n-2 (recursion), and returns the sum 
```

<!-- grade_id: cell-9b97ea3e85d77f4a, points: 0, grade: true, solution: false -->
```python
## tests ##

# Perform some tests    
assert f(0) == 0
assert f(1) == 1
assert f(2) == 1
assert f(3) == 2
assert f(10) == 55
assert f(15) == 610

# Check that ValueError is raised for n < 0
import pytest
with pytest.raises(ValueError):
    f(-1)
with pytest.raises(ValueError):
    f(-2)
```

## Exercise 09.2 (raising exceptions)

Modify your program from the bisection exercise in Activity 04 to raise an error if the maximum number of iterations is exceeded. Reduce the maximum allowed iterations to test that an exception is raised.

Add any other checks on the input data that you think are appropriate.

<!-- grade_id: cell-ca2c3ee62ad3b956, points: 0, grade: false, solution: true -->
```python
import pytest


def my_f(x):
    """Evaluate polynomial function"""
    return x*5 / 10 + x**3 - 10 * x*2 + 4 * x + 7

def compute_root(f, x0, x1, tol, max_it):
    # --- Input checks ---
    if not callable(f):
        raise ValueError("f must be a function.")
    if x0 >= x1:
        raise ValueError("x0 must be less than x1.")
    if tol <= 0:
        raise ValueError("Tolerance must be positive.")
    if max_it <= 0:
        raise ValueError("max_it must be positive.")

    f0 = f(x0)
    f1 = f(x1)

    # Check that the root is bracketed
    if f0 * f1 > 0:
        raise ValueError("Function has the same sign at x0 and x1 — no root guaranteed in this interval.")

    it = 0
    error = tol + 1
    x_mid = (x0 + x1) / 2

    while abs(error) > tol:
        x_mid = (x0 + x1) / 2
        f_mid = f(x_mid)
        print(f"Iteration {it}: x_mid={x_mid:.6f}, f_mid={f_mid:.6e}")

        if f0 * f_mid < 0:
            x1 = x_mid
            f1 = f_mid
        else:
            x0 = x_mid
            f0 = f_mid

        error = f_mid
        it += 1

        if it > max_it:
            raise RuntimeError(f"Maximum iterations ({max_it}) exceeded before convergence.")

    return x_mid, f_mid, it

```

<!-- grade_id: cell-f6cf91ed5dd1736b, points: 0, grade: true, solution: false -->
```python
## tests ##

# Test with max_it = 30
x, f, num_it = compute_root(my_f, x0=0, x1=2, tol=1.0e-6, max_it=30)

# Test with max_it = 20
with pytest.raises(RuntimeError):
    x, f, num_it = compute_root(my_f, x0=0, x1=2, tol=1.0e-6, max_it=20)
```

Output:
```
Iteration 0: x_mid=1.000000, f_mid=-7.500000e+00
Iteration 1: x_mid=0.500000, f_mid=-6.250000e-01
Iteration 2: x_mid=0.250000, f_mid=3.140625e+00
Iteration 3: x_mid=0.375000, f_mid=1.240234e+00
Iteration 4: x_mid=0.437500, f_mid=3.024902e-01
Iteration 5: x_mid=0.468750, f_mid=-1.626282e-01
Iteration 6: x_mid=0.453125, f_mid=6.959915e-02
Iteration 7: x_mid=0.460938, f_mid=-4.659891e-02
Iteration 8: x_mid=0.457031, f_mid=1.147920e-02
Iteration 9: x_mid=0.458984, f_mid=-1.756511e-02
Iteration 10: x_mid=0.458008, f_mid=-3.044265e-03
Iteration 11: x_mid=0.457520, f_mid=4.217140e-03
Iteration 12: x_mid=0.457764, f_mid=5.863553e-04
Iteration 13: x_mid=0.457886, f_mid=-1.228975e-03
Iteration 14: x_mid=0.457825, f_mid=-3.213152e-04
Iteration 15: x_mid=0.457794, f_mid=1.325187e-04
Iteration 16: x_mid=0.457809, f_mid=-9.439857e-05
Iteration 17: x_mid=0.457802, f_mid=1.906000e-05
Iteration 18: x_mid=0.457806, f_mid=-3.766930e-05
Iteration 19: x_mid=0.457804, f_mid=-9.304653e-06
Iteration 20: x_mid
```

```python

```